<a href="https://colab.research.google.com/github/harrisb002/CS_479/blob/Project-2/Filtering.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Non-linear filters

```python
import scipy
import scipy.ndimage as ndi
import numpy as np
import skimage
from skimage import io, filters
from skimage import data, img_as_float
import matplotlib.pyplot as plt
import os
import math
```

## Removing Gaussian Noise
A strategy to remove Gaussian noise is provided below. This is applicable when several images of the same object can be taken such as in astronomy or in microscopy. Due to sensor noise, each image is corrupted by Gaussian noise. In such cases, the algorithm below works well.

![Cleaning Gaussian Noise](cleaningGaussianNoise.jpeg)

## Part A (30 points)
Apply the algorithm given above to remove Gaussian noise from images in the directory `noisyImages` (already on the current path). It will require you to iterate over all images. Use `glob.glob(dir)` to extract list of `files` in `dir`. To store all images in a numpy array, it may be helpful to extract shape of any one image in the directory to be able to specify the overall size of the numpy array. Don't forget to convert your images to `float`!

<!-- grade_id: cell-69ac70b9df7d2608, points: 0, grade: false, solution: true -->
```python
import glob

def removeGaussianNoise(imgDir = 'wfc3_uvis'):

    dirname = os.path.join(os.getcwd(),'noisyImages',imgDir)
    numfiles = len(glob.glob(dirname + "/*.png"))
    print(dirname, 'Number of files = ' + str(numfiles), sep = '\n')


    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
from numpy.random import default_rng

#Helper function to display some images randomly from the directory imgDir,
#along with the clean image obtained by calling removeGaussianNoise
def showImages(imgDir, cleanImg):
    rng = default_rng()

    #Select some images randomly from imgDir and show them along with cleanImg
    dirname = os.path.join(os.getcwd(),'noisyImages',imgDir)
    numfiles = len(glob.glob(dirname + "/*.png"))
    files = []
    for file in glob.glob(dirname + "/*.png"):
        files.append(file)

    x = rng.choice(numfiles, size=3, replace=False)
    print("random indices:", x)
    fig, ax = plt.subplots(nrows=2, ncols=2, figsize=(15, 10), sharex = True, sharey = True)
    ax = ax.ravel()
    for i in range(3):
        img = io.imread(files[x[i]])
        ax[i].imshow(img)

    ax[3].imshow(cleanImg)
    plt.tight_layout()

```

<!-- grade_id: cell-29feaa97827ad57d, points: 10, grade: true, solution: false -->
```python
imgDir = 'wfc3_uvis'
cleanImg = removeGaussianNoise(imgDir)
showImages(imgDir, cleanImg)

assert cleanImg.shape[0] == 346
assert math.isclose(cleanImg.mean(), 0.25733878625)

print("Success!")
```

<!-- grade_id: cell-994472389df1d9df, points: 10, grade: true, solution: false -->
```python
imgDir = 'orion'
cleanImg = removeGaussianNoise(imgDir)

showImages(imgDir, cleanImg)

assert cleanImg.shape[0] == 600
assert math.isclose(cleanImg.mean(), 0.23500850036)

print("Success!")
```

```python
imgDir = 'n44f'
cleanImg = removeGaussianNoise(imgDir)

showImages(imgDir, cleanImg)
print(cleanImg.mean())
assert cleanImg.shape[0] == 339
assert math.isclose(cleanImg.mean(), 0.2849458358)

print("Success!")
```

## Part C (20 points)

Check out the description of the problem of Pesky Tourist [here](http://nifty.stanford.edu/2014/nicholson-the-pesky-tourist/student.html). You may see the images in the directory `ThePeskyTourist` provided as a part of this project. Your task is to remove the Pesky Tourist from the images! Ghostly tourist is not acceptable :) The algorithm is very similar to above (that removes Gaussian noise) but taking a median instead of taking an average.

**Note**: do NOT change the location of the directory `ThePeskyTourist` relative to this notebook. Otherwise the autograder will call foul.

<!-- grade_id: cell-a2ae3a8b94bc5603, points: 0, grade: false, solution: true -->
```python
def removePeskyTourist():

    # YOUR CODE HERE
    raise NotImplementedError()
```

Let us visualize some random images of the tourist along with the cleaned version

```python
from numpy.random import default_rng

rng = default_rng()

#Find the clean version
cleanImg = removePeskyTourist()

#Select some images randomly from imgDir and show them along with cleanImg
dirname = os.path.join(os.getcwd(),'ThePeskyTourist')
numfiles = len(glob.glob(dirname + "/*.png"))
files = []
for file in glob.glob(dirname + "/*.png"):
    files.append(file)

x = rng.choice(numfiles, size=3, replace=False)
print(x)
fig, ax = plt.subplots(nrows=2, ncols=2, figsize=(15, 10), sharex = True, sharey = True)
ax = ax.ravel()
for i in range(3):
    img = io.imread(files[x[i]])
    ax[i].imshow(img)

ax[3].imshow(cleanImg)
plt.tight_layout()

```

<!-- grade_id: cell-0206042a9e5412e4, points: 20, grade: true, solution: false -->
```python
cleanImg = removePeskyTourist()

assert cleanImg.shape[0] == 557
assert math.isclose(cleanImg.mean(), 0.52157602355)

print("Success!")


```

## Part C: Coin Challenge (DEMO ONLY)

Load the coin image. Create a binary image (mask) by thresholding using Otsu Threshold. Process the mask below using morphological operations to label the 24 coins separately. Your code must show there are exactly 24 labels (using region properties). It is not necessary to create the binary image using Otsu Threshold. As long as you are able to segment the 24 coins faithfully, you will get full credit.

Example output:

![output](coins_challenge.png)

<!-- grade_id: cell-db60165f5e2391e4, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```